# Lab 5 - Getting Data

The topics of week 5 is getting data with web scraping. 

In this lab notebook you will gain experience webscraping data of various forms.

We will be looking at various US government website and using them to get information (these sites are selected because of their public informative nature). 

## Autograder Setup

The next code cell should be uncommented to run the autograder tests when using Colab/DeepNote.  If you are using an environment with `otter-grader` already installed (your own machine, lab machines), then do not uncomment the code.

```python
# !pip install otter-grader
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Collecting otter-grader
  Downloading otter_grader-4.0.2-py3-none-any.whl (167 kB)
[K     |████████████████████████████████| 167 kB 11.3 MB/s 
[?25hRequirement already satisfied: wrapt in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.14.1)
Requirement already satisfied: click in /usr/local/lib/python3.7/dist-packages (from otter-grader) (7.1.2)
Requirement already satisfied: google-api-python-client in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.12.11)
Collecting fica>=0.2.0
  Downloading fica-0.2.0-py3-none-any.whl (11 kB)
Requirement already satisfied: nbformat in /usr/local/lib/python3.7/dist-packages (from otter-grader) (5.4.0)
Requirement already satisfied: pandas in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.3.5)
Requirement already satisfied: jinja2 in /usr/local/lib/python3.7/dist-packages (from otter-grader) (2.11.3)
Requirem
```

In order to have the data files and test files for this lab, uncomment and run the following cell. Please comment out before submission.

```python
# !wget https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab5/lab5.files.zip
# !unzip lab5.files.zip
```

Output:
```
--2022-09-29 22:09:06--  https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab5/lab5.files.zip
Resolving pages.mtu.edu (pages.mtu.edu)... 141.219.70.232
Connecting to pages.mtu.edu (pages.mtu.edu)|141.219.70.232|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 2016 (2.0K) [application/zip]
Saving to: ‘lab5.files.zip’

lab5.files.zip      100%[===================>]   1.97K  --.-KB/s    in 0s      

2022-09-29 22:09:06 (282 MB/s) - ‘lab5.files.zip’ saved [2016/2016]

Archive:  lab5.files.zip
   creating: tests/
  inflating: tests/q5.py             
  inflating: tests/q1.py             
  inflating: tests/q4.py             
  inflating: tests/q3.py             
  inflating: tests/q2.py
```

## Lab Setup 

```python
from bs4 import BeautifulSoup
import requests 
import pandas as pd
import re

import otter
grader = otter.Notebook()
```

## Exercise 1 

Identify the name of the data set most relevant to 'climate' on https://www.data.gov 

You will want to use the url = "https://catalog.data.gov/dataset?q=climate&sort=score+desc%2C+name+asc".

Then, you will use BeautifulSoup to parse the website to find the name of the data set.  

```python
# Identify the name of the data set most relevant to "climate" on data.gov
# Use BeautifulSoup to parse the website to find the name of the data set. 
url = 'https://catalog.data.gov/dataset?q=climate&sort=score+desc%2C+name+asc'
resp = requests.get(url) 
soup = BeautifulSoup(resp.text, 'html5lib')

name = soup.find('h3').text.strip()
name
```

Output:
```
'Downscaled Climate Model Climate Toolbox'
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Exercise 2 

Total number of clinical trials (research studies) recorded at the National Institutes of Health, https://clinicaltrials.gov

```python
# Identify the number of clinical trials (research studies) on https://clinicaltrials.gov
# Use BeautifulSoup to parse the website to find the number of studies. 
url = 'https://clinicaltrials.gov/'
resp = requests.get(url) 
soup = BeautifulSoup(resp.text, 'html5lib')

pattern = r'\d+?,\d+'

num_research_studies = int(re.search(pattern, soup.find('h2').text).group().replace(',', ''))  # integer value
num_research_studies
```

Output:
```
429205
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Exercise 3 

Give the number of travel advisories, by level, for international travel given by the US government. Data available at: https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/

Create two solutions for this question, the first using `pandas` `read_html` function.  The second version should be using just BeautifulSoup and python to scrape the table.  

```python
site = requests.get('https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/').text
q3 = BeautifulSoup(site, 'html5lib')
```

**Version 1** - Write one solution using pandas `read_html`

**Version 2** - Write another solution just using beautiful soup and other Python functions to scrape the table.

```python
# Version 1
# Read in the table using pandas "read_html" function 
# Report out number of each warnings 
...
url = 'https://travel.state.gov/content/travel/en/traveladvisories/traveladvisories.html/'
df = pd.read_html(url)[0]

l1Num1 = df['Level'].value_counts()[0] 
l2Num1 = df['Level'].value_counts()[1]
l3Num1 = df['Level'].value_counts()[2]
l4Num1 = df['Level'].value_counts()[3]

print("Number of travel Level 1 warnings : " +  str(l1Num1)) 
print("Number of travel Level 2 warnings : " +  str(l2Num1))
print("Number of travel Level 3 warnings : " +  str(l3Num1))
print("Number of travel Level 4 warnings : " +  str(l4Num1))
```

Output:
```
Number of travel Level 1 warnings : 89
Number of travel Level 2 warnings : 62
Number of travel Level 3 warnings : 39
Number of travel Level 4 warnings : 20
```

```python
# Version 2
# Write another solution just using beautiful soup and other Python functions to 
#   scrape the table information
# Report out number of each warnings 
...

all_enries = q3.find_all('td')
levels = [0,0,0,0]
pattern = r'Level (\d)'
#re.match(pattern, all_enries[1].text).group(1)

for entry in all_enries:
  m = re.match(pattern, entry.text)
  if m:
    level = int(m.group(1))
    levels[level-1] += 1


l1Num2 = levels[0]
l2Num2 = levels[1]
l3Num2 = levels[2]
l4Num2 = levels[3]

print("Number of travel Level 1 warnings : " +  str(l1Num2)) 
print("Number of travel Level 2 warnings : " +  str(l2Num2))
print("Number of travel Level 3 warnings : " +  str(l3Num2))
print("Number of travel Level 4 warnings : " +  str(l4Num2))
```

Output:
```
Number of travel Level 1 warnings : 89
Number of travel Level 2 warnings : 62
Number of travel Level 3 warnings : 39
Number of travel Level 4 warnings : 20
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

##  Example 

Let's look at the number of workplace fatalities as reported to the federal and state Occupational Safety and Health Administration (OSHA) at https://www.osha.gov/fatalities. 




```python
site = requests.get('https://www.osha.gov/fatalities').text
q4 = BeautifulSoup(site, 'html5lib')

# Read in Table using pandas "read_html" method
q4table = pd.read_html(q4.prettify(), flavor=['bs4'])
df4 = q4table[0]
#df4.head()
```

```python
df4.shape
```

Output:
```
(6252, 7)
```

```python
sz = df4[df4['Citation Issued Related to Fatality']=='Yes'].shape
print("Number of citations issues related to fatality: ", sz[0])
```

Output:
```
Number of citations issues related to fatality:  3446
```

## Exercise 4 

Save in a DataFrame the incidents that occurred in fiscal year 2021, Oct. 1, 2020 - Sept. 30, 2021.  

Try using `pd.to_datetime` to add a new "Date" column with the Date of Incident column converted to a datetime object that may be easier to filter. 

https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html


```python
# Save all incidents that occured in fiscal year 2021: Oct. 1, 2020 - Sept. 30, 2021
#  in a DataFrame ("q4result"), it will have a new column "Date" with the Date of 
#  Incident column converted to a datetime object. 

df4['Date'] = pd.to_datetime(df4['Date of Incident'])
q4result = df4.loc[ (df4['Date']>=pd.to_datetime('2020-10-01')) & (df4['Date']<=pd.to_datetime('2021-09-30')) ]
q4result
```

Output:
```
Date of Incident        City Select State  \
287        09/28/2021   Stoughton           WI   
288        09/28/2021   Lindstrom           MN   
289        09/25/2021    Clermont           FL   
290        09/24/2021     Kooskia           ID   
291        09/24/2021   Charlotte           NC   
...               ...         ...          ...   
1400       10/02/2020  Germantown           NY   
1401       10/02/2020  Sacramento           CA   
1402       10/01/2020    Au Train           MI   
1403       10/01/2020  Petersburg           VA   
1404       10/01/2020     Waldorf           MD   

                                     Hazard Description  Inspection Number  \
287   Robert Castillo (87) fatally crushed in trench...          1551677.0   
288     David Chmielewski (48) died in trench collapse.          1560134.0   
289   Jesus Olvera Aguilar (19) fatally struck by li...          1551830.0   
290               Woker died in fall from trailer deck.          1554698.0   
291   Doyle "G
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Exercise 5 

Identify the most viewed data set on Michigan state's open data portal.   
https://data.michigan.gov/browse?sortBy=most_accessed

Stores its name in the provided "most_viewed_data" variable and the number of views in "num_views".

```python
#browse2-result-view-count-value

url = 'https://data.michigan.gov/browse?sortBy=most_accessed'
resp = requests.get(url) 
soup = BeautifulSoup(resp.text, 'html5lib')



most_viewed_data = soup.find('h2', class_= 'browse2-result-name').text.strip()
num_views = int(soup.find('div', class_='browse2-result-view-count-value').text.strip().replace(',',''))

print(most_viewed_data)
print(num_views)
```

Output:
```
MI_Congress_Plot
44291
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

Tests for q1 (points: 1)
```
>>> name == "Downscaled Climate Model Climate Toolbox"
True
```


Tests for q2 (points: 1)
```
>>> num_research_studies >= 429205 and num_research_studies <= 429205+1600
True
```


Tests for q3 (points: 2)
```
>>> l1Num1 == l1Num2 
True
```
```
>>> l2Num1 == l2Num2 
True
```
```
>>> l3Num1 == l3Num2
True
```
```
>>> l4Num1 == l4Num2
True
```
```
>>> l1Num1 == 89
True
```
```
>>> l2Num1 == 62
True
```
```
>>> l3Num1 == 39 
True
```
```
>>> l4Num1 == 20
True
```


Tests for q4 (points: 2)
```
>>> q4result.shape[0] == 1118 
True
```
```
>>> q4result.shape[1] == 8
True
```
```
>>> "Date" in q4result.columns
True
```
```
>>> int(q4result['Inspection Number'].min()) == 1495370
True
```
```
>>> int(q4result['Inspection Number'].max()) == 1580910
True
```


Tests for q5 (points: 2)
```
>>> most_viewed_data == 'MI_Congress_Plot'
True
```
```
>>> num_views >= 44271 and num_views <= 44271+5000
True
```
