# Lab 7 - Relational Algebra 

This examples build on examples from Data100 and also Ch. 23 of the Data Science by Scratch book. 

In this lab notebook you will write functions to implement the basic relational algebra operators.

## Autograder Setup 

The next code cell should be uncommented to run the autograder tests when using Colab/DeepNote. If you are using an environment with otter-grader already installed (your own machine, lab machines), then do not uncomment the code.

```python
# !pip install otter-grader
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Requirement already satisfied: otter-grader in /usr/local/lib/python3.7/dist-packages (4.1.0)
Requirement already satisfied: requests in /usr/local/lib/python3.7/dist-packages (from otter-grader) (2.23.0)
Requirement already satisfied: gspread in /usr/local/lib/python3.7/dist-packages (from otter-grader) (3.4.2)
Requirement already satisfied: jinja2 in /usr/local/lib/python3.7/dist-packages (from otter-grader) (2.11.3)
Requirement already satisfied: pandas in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.3.5)
Requirement already satisfied: click in /usr/local/lib/python3.7/dist-packages (from otter-grader) (7.1.2)
Requirement already satisfied: PyYAML in /usr/local/lib/python3.7/dist-packages (from otter-grader) (6.0)
Requirement already satisfied: wrapt in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.14.1)
Requirement already satisfied: six in /usr/local/
```

In order to have the data files and test files for this lab, uncomment and run the following cell. 

**Please comment out these two cells before submission.**

```python
# !wget https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab7/lab7.files.zip
# !unzip lab7.files.zip
```

Output:
```
--2022-10-14 02:15:47--  https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab7/lab7.files.zip
Resolving pages.mtu.edu (pages.mtu.edu)... 141.219.70.232
Connecting to pages.mtu.edu (pages.mtu.edu)|141.219.70.232|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 3168 (3.1K) [application/zip]
Saving to: ‘lab7.files.zip.3’

lab7.files.zip.3    100%[===================>]   3.09K  --.-KB/s    in 0s      

2022-10-14 02:15:47 (499 MB/s) - ‘lab7.files.zip.3’ saved [3168/3168]

Archive:  lab7.files.zip
replace tests/q5.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q1.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q4.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q8.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q3.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q7.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q6.py? [y]es, [n]o, [A]ll, [N]one, [r]ename: n
replace tests/q2.py? [y]es, [n]o, [A]ll, [N]one, [r]
```

*Remember that you answers will potentially be evaluated on additional hidden test cases as well as manual code review.* 

## Lab Setup

```python
import numpy as np
import pandas as pd

import otter
grader = otter.Notebook()
```

# Huskies Sailing Club 

The Huskies Sailing Club wants to better organize their user data. Your first job is to implement code relational algebra operators in Python. 

You may want to refer to the [pandas](http://pandas.pydata.org/pandas-docs/stable/) documentation. 

Here are the Sailing Club's databases.  Your job implement a variety of unary and binary relational algebra operators. 

```python
young_sailors = pd.DataFrame({
        "sid":      [2701, 18869, 63940, 21869, 17436],
        "sname":    ["Grace", "Morgan", "Anne", "Jack", "Dustin"],
        "rating":   [8, 6, 4, 9, 3],
        "age":      [25, 26, 21, 27, 22],
    })
salty_sailors = pd.DataFrame({
        "sid":      [2701, 17436, 45433, 22689, 46535],
        "sname":    ["Grace", "Dustin", "Balon", "Euron", "Victarion"],
        "rating":   [8, 3, 7, 10, 2],
        "age":      [25, 22, 39, 35, 37],
    })
boats = pd.DataFrame({
        "bid":      [41116, 54505, 50041, 35168, 58324],
        "bname":    ["The Black Sparrow", "The Great Kraken", "The Prophetess", "Silence", "Iron Victory"],
        "color":    ["Black", "Orange", "Silver", "Red", "Grey"],
    })
reservations = pd.DataFrame({
        "sid":      [21869, 45433, 18869, 22689, 21869, 17436, 63940, 45433, 21869, 18869],
        "bid":      [41116, 35168, 50041, 41116, 58324, 50041, 54505, 41116, 50041, 41116],
        "day":      ["3/1", "3/1", "3/2", "3/2", "3/2", "3/3", "3/3", "3/3", "3/3", "3/4"],
    })
```

```python
young_sailors
```

Output:
```
sid   sname  rating  age
0   2701   Grace       8   25
1  18869  Morgan       6   26
2  63940    Anne       4   21
3  21869    Jack       9   27
4  17436  Dustin       3   22
```

```python
salty_sailors
```

Output:
```
sid      sname  rating  age
0   2701      Grace       8   25
1  17436     Dustin       3   22
2  45433      Balon       7   39
3  22689      Euron      10   35
4  46535  Victarion       2   37
```

```python
boats
```

Output:
```
bid              bname   color
0  41116  The Black Sparrow   Black
1  54505   The Great Kraken  Orange
2  50041     The Prophetess  Silver
3  35168            Silence     Red
4  58324       Iron Victory    Grey
```

```python
reservations
```

Output:
```
sid    bid  day
0  21869  41116  3/1
1  45433  35168  3/1
2  18869  50041  3/2
3  22689  41116  3/2
4  21869  58324  3/2
5  17436  50041  3/3
6  63940  54505  3/3
7  45433  41116  3/3
8  21869  50041  3/3
9  18869  41116  3/4
```

## Exercise 1: Projection

Create a function `project` that takes as arguments a dataframe and a list of columns to select and returns the projection.

```python
def project(df, columns):     
    return pd.DataFrame(df, columns=columns)

df1a = project(salty_sailors, ["sname", "age"])
df1b = project(boats, ["color"])
```

Example test:  
The output for `project(salty_sailors, ["sname", "age"])` should be: 

|   | sname  |  age  | 
|---|--------|-------|
| 0 | Grace  | 25    |
| 1 | Dustin | 22    | 
| 2 | Balon  | 39    | 
| 3 | Euron  | 35    | 
| 4 | Victarion | 37 |


```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Exercise 2: Selection 

Create a function `select` that has arguments of a data frame and a condition (function) which determines which rows to select. 

For example: 
`good_sailors = select(young_sailors, lambda x: x["rating"] > 6)`

```python
def select(df, condition):     
    return df[condition]

dfq2 = select(young_sailors, lambda x: x["rating"] > 6)
```

The output for `select(young_sailors, lambda x: x["rating"] > 6)` should be:

|   | sid   | sname  | rating | age |
|---|-------|--------|--------|-----|
| 0 | 2701  | Grace  | 8      | 25  |
| 3 | 21869 | Jack   | 9      | 27  |

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Exercise 3: Union 

Create a function, `union`, (a binary operator) that takes two dataframes as input.  You can assume the two dataframes are union compatible (that is, they have the same number of columns, and their columns have the same types).  Return the result of the union operation.  Assume set semantics, that is, duplicates should be removed. 

```python
def union(df1, df2):     
    return pd.concat([df1, df2]).drop_duplicates()

df3 = union(young_sailors, salty_sailors)
df3
```

Output:
```
sid      sname  rating  age
0   2701      Grace       8   25
1  18869     Morgan       6   26
2  63940       Anne       4   21
3  21869       Jack       9   27
4  17436     Dustin       3   22
2  45433      Balon       7   39
3  22689      Euron      10   35
4  46535  Victarion       2   37
```

The output for `union(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 0 | 2701  | Grace    | 8      | 25  |
| 1 | 18869 | Morgan   | 6      | 26  | 
| 2 | 63940 | Anne     | 4      | 21  | 
| 3 | 21869 | Jack     | 9      | 27  | 
| 4 | 17436 | Dustin   | 3      | 22  | 
| 2 | 45433 | Balon    | 7      | 39  | 
| 3 | 22689 | Euron    | 10     | 35  | 
| 4 | 46535 | Victarion | 2     | 37  | 


```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Exercise 4: Intersection 

Write a function to perform the binary operation of Intersection.

```python
def intersection(df1, df2):     
    return pd.merge(df1, df2)

df4 = intersection(young_sailors, salty_sailors)
```

The output for `intersection(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 0 | 2701  | Grace    | 8      | 25  |
| 1 | 17436 | Dustin   | 3      | 22  | 


```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Exercise 5: Set Difference 

Write a function to perform the set difference operation.  This one is a bit harder.  You may want to convert rows of the dataframe to a tuple, if you are having trouble. Alternatively, you could consider using the `isin` method.

```python
def difference(df1, df2): 
    # Assume inputs df1 and df2 are same schema 
    # Make sure your returned DataFrame has the same data types as the 
    #  input DataFrames 
    # Make sure to not change the original data frames passed as 
    #  inputs
    df = pd.concat([pd.merge(df1, df2, how="left"), pd.merge(df1, df2, how="inner")])
    return df.drop_duplicates(keep=False)

df5 = difference(young_sailors, salty_sailors)
df5
```

Output:
```
sid   sname  rating  age
1  18869  Morgan       6   26
2  63940    Anne       4   21
3  21869    Jack       9   27
```

The output for `difference(young_sailors, salty_sailors)` should be:

|   | sid   | sname    | rating | age |
|---|-------|----------|--------|-----|
| 1 | 18869.0 | Morgan   | 6.0      | 26.0  | 
| 2 | 63940.0 | Anne     | 4.0      | 21.0  | 
| 3 | 21869.0 | Jack     | 9.0      | 27.0  | 


```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

## Exercise 6: Cross-product 

This one is also harder, so there is some code to start.  Think about how the new `tmp-key` column could be used ...

```python
def cross_product(in1, in2):
    # add a column "tmp-key" of zeros to df1 and df2 
    in1 = pd.concat([in1, pd.Series(0, index=in1.index, name="tmp-key")], axis=1)
    in2 = pd.concat([in2, pd.Series(0, index=in2.index, name="tmp-key")], axis=1)
    # use Pandas merge functionality along with drop 
    # to compute outer product and remove extra column
    return pd.merge(in1, in2, how="cross").drop(['tmp-key_x', 'tmp-key_y'], axis =1)

df6 = cross_product(young_sailors, salty_sailors)
```

The output of `cross_product(young_sailors, salty_sailors)` should be: 

|   | sid_x | sname_x  | rating_x | age_x | sid_y  | sname_y | rating_y | age_y |
|---|-------|----------|--------|-----|-------|--------|-----|----|
| 0 | 2701  | Grace    | 8      | 25  | 2701  | Grace  | 8   | 25 |
| 1 | 2701  | Grace    | 8      | 25  | 17436 | Dustin | 3   | 22 |
| 2 | 2701  | Grace    | 8      | 25  | 45433 | Balon  | 7   | 39 | 
| 3 | 2701  | Grace    | 8      | 25  | 22689 | Euron  | 10  | 35 |
| 4 | 2701  | Grace    | 8      | 25  | 46535 | Victarion | 2   | 37 |
| 5 | 18869 | Morgan   | 6      | 26  | 2701  | Grace  | 8   | 25 |
|   | ...   | ...      | ...    | ... | ...   | ...    | ... | ... |

There should be 25 entries.

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## Exercise 7: Theta-Join 

Complete a function for theta-join.  Consider using two other relational operators. 

```python
def theta_join(df1, df2, condition): 
  return select(pd.merge(df1, df2, how="cross"), condition)
df7 = theta_join(young_sailors, salty_sailors, lambda x: x["age_x"] > x["age_y"])
```

The output of `theta_join(young_sailors, salty_sailors, lambda x: x["age_x"] > x["age_y"])` should be: 

|   | sid_x | sname_x  | rating_x | age_x | sid_y  | sname_y | rating_y | age_y |
|---|-------|----------|--------|-----|-------|--------|-----|----|
| 0 | 2701  | Grace    | 8      | 25  | 17436 | Dustin | 3   | 22 |
| 5 | 18869 | Morgan   | 6      | 26  | 2701  | Grace  | 8   | 25 |
| 6 | 18869 | Morgan   | 6      | 26  | 17436 | Dustin | 3   | 22 | 
| 15| 21869 | Jack     | 9      | 27  | 2701  | Grace  | 8   | 25 |
| 16| 21869 | Jack     | 9      | 27  | 17436 | Dustin | 3   | 22 |


```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## Exercise 8: Natural Join 

Implement a function to return the natural join.  Try to implement this using other relational operators.

```python
def natural_join(df1, df2, attr):
  df = df1.merge(df2, how='cross')
  return df[df['sid_x'] == df['sid_y']]

all_sailors = union(young_sailors, salty_sailors)
df8 = natural_join(all_sailors, reservations, "sid")
print(df8)
```

Output:
```
sid_x   sname  rating  age  sid_y    bid  day
12  18869  Morgan       6   26  18869  50041  3/2
19  18869  Morgan       6   26  18869  41116  3/4
26  63940    Anne       4   21  63940  54505  3/3
30  21869    Jack       9   27  21869  41116  3/1
34  21869    Jack       9   27  21869  58324  3/2
38  21869    Jack       9   27  21869  50041  3/3
45  17436  Dustin       3   22  17436  50041  3/3
51  45433   Balon       7   39  45433  35168  3/1
57  45433   Balon       7   39  45433  41116  3/3
63  22689   Euron      10   35  22689  41116  3/2
```

The output of the natural join is: 

|   | sid_x |	sname |	rating | age  | sid_y | bid | day |
|---|-------|---------|--------|------|-------|-----|-----|
| 12 |	18869 |	Morgan | 6 | 26 | 18869 | 50041 | 3/2 |
| 19 |	18869 |	Morgan | 6 | 26 | 18869 | 41116 | 3/4 |
| 26 |	63940 |	Anne | 4 | 21 |	63940 |	54505 |	3/3 |
| 30 |	21869 |	Jack | 9 | 27 |	21869 |	41116 |	3/1 |
| 34 |	21869 |	Jack | 9 | 27 |	21869 |	58324 |	3/2 |
| 38 |	21869 |	Jack | 9 | 27 |	21869 |	50041 |	3/3 |
| 45 |	17436 |	Dustin | 3 | 22 |	17436 |	50041 |	3/3 |
| 51 |	45433 |	Balon |	7 |	39 |	45433 |	35168 |	3/1 |
| 57 |	45433 |	Balon |	7 |	39 |	45433 |	41116 |	3/3 |
| 63 |	22689 |	Euron |	10 | 35 |	22689 |	41116 |	3/2 |

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

Tests for q1 (points: 1)
```
>>> df1a.shape[0]==5 and df1a.shape[1]==2 
True
```
```
>>> all(df1a.columns == ["sname", "age"])
True
```
```
>>> df1a.iloc[1,1]==22 and df1a.iloc[3,0]=='Euron'
True
```
```
>>> df1b.shape[0]==5 and df1b.shape[1]==1
True
```
```
>>> all(df1b.columns=='color') and df1b.iloc[2,0]=='Silver'
True
```


Tests for q2 (points: 1)
```
>>> dfq2.shape[0]==2 and dfq2.shape[1]==4
True
```
```
>>> dfq2.iloc[1,3]==27 and dfq2.iloc[0,2]==8
True
```


Tests for q3 (points: 1)
```
>>> df3.shape[0]==8 and df3.shape[1]==4
True
```


Tests for q4 (points: 1)
```
>>> intersection(young_sailors, salty_sailors).shape == (2,4)
True
```


Tests for q5 (points: 2)
```
>>> df5.shape == (3,4)
True
```
```
>>> df5.iloc[2,:].to_list() == [21869, 'Jack', 9, 27]
True
```
```
>>> df5.dtypes[0] == 'int64'
True
```
```
>>> df5.dtypes[3] == 'int64'
True
```


Tests for q6 (points: 2)
```
>>> df6.shape == (25,8)
True
```
```
>>> df6.iloc[8,3:6].to_list() == [26, 22689, 'Euron']
True
```


Tests for q7 (points: 1)
```
>>> df7.shape == (5,8)
True
```
```
>>> df7.iloc[3,2:5].to_list() == [9, 27, 2701]
True
```


Tests for q8 (points: 1)
```
>>> df8.shape == (10,7)
True
```
```
>>> df8.iloc[5,1:4].to_list() == ['Jack', 9, 27]
True
```
