# Getting Data
The topics of week 6 continues to be getting data, in this case using an API to access structured data. 

In this lab notebook you will gain experience reading data from and posting to an API. 


## Autograder Setup 

The next code cell should be uncommented to run the autograder tests when using Colab/DeepNote. If you are using an environment with otter-grader already installed (your own machine, lab machines), then do not uncomment the code.

```python
#!pip install otter-grader
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Requirement already satisfied: otter-grader in /usr/local/lib/python3.7/dist-packages (4.1.0)
Requirement already satisfied: google-api-python-client in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.12.11)
Requirement already satisfied: gspread in /usr/local/lib/python3.7/dist-packages (from otter-grader) (3.4.2)
Requirement already satisfied: wrapt in /usr/local/lib/python3.7/dist-packages (from otter-grader) (1.14.1)
Requirement already satisfied: PyYAML in /usr/local/lib/python3.7/dist-packages (from otter-grader) (6.0)
Requirement already satisfied: requests in /usr/local/lib/python3.7/dist-packages (from otter-grader) (2.23.0)
Requirement already satisfied: python-on-whales in /usr/local/lib/python3.7/dist-packages (from otter-grader) (0.52.0)
Requirement already satisfied: click in /usr/local/lib/python3.7/dist-packages (from otter-grader) (7.1.2)
Requirement already
```

In order to have the data files and test files for this lab, uncomment and run the following cell. Please comment out before submission.

```python
#!wget https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab6/lab6.files.zip
#!unzip lab6.files.zip
```

Output:
```
--2022-10-11 16:06:47--  https://pages.mtu.edu/~lebrown/un5550-f22/labs/lab6/lab6.files.zip
Resolving pages.mtu.edu (pages.mtu.edu)... 141.219.70.232
Connecting to pages.mtu.edu (pages.mtu.edu)|141.219.70.232|:443... connected.
HTTP request sent, awaiting response... 200 OK
Length: 831 [application/zip]
Saving to: ‘lab6.files.zip.1’

lab6.files.zip.1    100%[===================>]     831  --.-KB/s    in 0s      

2022-10-11 16:06:48 (111 MB/s) - ‘lab6.files.zip.1’ saved [831/831]

Archive:  lab6.files.zip
replace tests/q1.py? [y]es, [n]o, [A]ll, [N]one, [r]ename:
```

*Remember that you answers will potentially be evaluated on additional hidden test cases as well as manual code review.* 

## Lab Setup

```python
import requests
import json
import datetime
from io import StringIO
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import matplotlib as mpl 
%matplotlib inline  

import otter
grader = otter.Notebook()
```

## API Getting Data

So far we have seen examples of getting data from an API.  These examples make use of GET requests from the API/server. 

Making a HTTP GET request can be done using several python libraries including: 

* httplib 
* urllib 
* requests 

We have been using the `requests` module.

Let's look at another example.

## Example: Google Books

Here we will examine using the Google Books API:  
https://developers.google.com/books/docs/overview

```python
# api-endpoint 
url = "https://www.googleapis.com/books/v1/volumes"
  
isbn = "isbn:0553386794"

# set the parameters to be sent to the API
params = {'q': isbn}

resp = requests.get(url, params)
```

Look at what the response is? 

How do we then extract the data?

```python
resp
```

Output:
```
<Response [200]>
```

```python
dat = resp.json()
dat
```

Output:
```
{'kind': 'books#volumes',
 'totalItems': 2,
 'items': [{'kind': 'books#volume',
   'id': 'CKhFzQEACAAJ',
   'etag': '50b7MI1EgQo',
   'selfLink': 'https://www.googleapis.com/books/v1/volumes/CKhFzQEACAAJ',
   'volumeInfo': {'title': 'A Song of Ice and Fire',
    'authors': ['George R. R. Martin'],
    'publishedDate': '2011',
    'description': 'The uneasy peace that exists following the death of Robb Stark is threatened by new plots, intrigues, and alliances that once again will plunge the Seven Kingdoms back into all-out war for control of the Iron Throne.',
    'industryIdentifiers': [{'type': 'ISBN_10', 'identifier': '0553386794'},
     {'type': 'ISBN_13', 'identifier': '9780553386790'}],
    'readingModes': {'text': False, 'image': False},
    'printType': 'BOOK',
    'averageRating': 4,
    'ratingsCount': 18,
    'maturityRating': 'NOT_MATURE',
    'allowAnonLogging': False,
    'contentVersion': 'preview-1.0.0',
    'panelizationSummary': {'containsEpubBubbles': False,
     'co
```

There is a lot of information here.  Explore the structure of the JSON information. 

```python
# First, we can print it better! 
print(json.dumps(resp.json(), indent=2))
```

Output:
```
{
  "kind": "books#volumes",
  "totalItems": 2,
  "items": [
    {
      "kind": "books#volume",
      "id": "CKhFzQEACAAJ",
      "etag": "50b7MI1EgQo",
      "selfLink": "https://www.googleapis.com/books/v1/volumes/CKhFzQEACAAJ",
      "volumeInfo": {
        "title": "A Song of Ice and Fire",
        "authors": [
          "George R. R. Martin"
        ],
        "publishedDate": "2011",
        "description": "The uneasy peace that exists following the death of Robb Stark is threatened by new plots, intrigues, and alliances that once again will plunge the Seven Kingdoms back into all-out war for control of the Iron Throne.",
        "industryIdentifiers": [
          {
            "type": "ISBN_10",
            "identifier": "0553386794"
          },
          {
            "type": "ISBN_13",
            "identifier": "9780553386790"
          }
        ],
        "readingModes": {
          "text": false,
          "image": false
        },
        "printType": "BOOK",
        "av
```

```python
dat.keys()
```

Output:
```
dict_keys(['kind', 'totalItems', 'items'])
```

```python
dat['kind']
```

Output:
```
'books#volumes'
```

```python
dat['totalItems']
```

Output:
```
2
```

```python
dat['items']
```

Output:
```
[{'kind': 'books#volume',
  'id': 'CKhFzQEACAAJ',
  'etag': '50b7MI1EgQo',
  'selfLink': 'https://www.googleapis.com/books/v1/volumes/CKhFzQEACAAJ',
  'volumeInfo': {'title': 'A Song of Ice and Fire',
   'authors': ['George R. R. Martin'],
   'publishedDate': '2011',
   'description': 'The uneasy peace that exists following the death of Robb Stark is threatened by new plots, intrigues, and alliances that once again will plunge the Seven Kingdoms back into all-out war for control of the Iron Throne.',
   'industryIdentifiers': [{'type': 'ISBN_10', 'identifier': '0553386794'},
    {'type': 'ISBN_13', 'identifier': '9780553386790'}],
   'readingModes': {'text': False, 'image': False},
   'printType': 'BOOK',
   'averageRating': 4,
   'ratingsCount': 18,
   'maturityRating': 'NOT_MATURE',
   'allowAnonLogging': False,
   'contentVersion': 'preview-1.0.0',
   'panelizationSummary': {'containsEpubBubbles': False,
    'containsImageBubbles': False},
   'language': 'en',
   'previewLink': 'htt
```

`dat['items']` returns a list of items.

```python
type(dat['items'])
```

Output:
```
list
```

```python
# We can look at the first item on the list 
dat['items'][0]
```

Output:
```
{'kind': 'books#volume',
 'id': 'CKhFzQEACAAJ',
 'etag': '50b7MI1EgQo',
 'selfLink': 'https://www.googleapis.com/books/v1/volumes/CKhFzQEACAAJ',
 'volumeInfo': {'title': 'A Song of Ice and Fire',
  'authors': ['George R. R. Martin'],
  'publishedDate': '2011',
  'description': 'The uneasy peace that exists following the death of Robb Stark is threatened by new plots, intrigues, and alliances that once again will plunge the Seven Kingdoms back into all-out war for control of the Iron Throne.',
  'industryIdentifiers': [{'type': 'ISBN_10', 'identifier': '0553386794'},
   {'type': 'ISBN_13', 'identifier': '9780553386790'}],
  'readingModes': {'text': False, 'image': False},
  'printType': 'BOOK',
  'averageRating': 4,
  'ratingsCount': 18,
  'maturityRating': 'NOT_MATURE',
  'allowAnonLogging': False,
  'contentVersion': 'preview-1.0.0',
  'panelizationSummary': {'containsEpubBubbles': False,
   'containsImageBubbles': False},
  'language': 'en',
  'previewLink': 'http://books.google.nl/b
```

```python
'''We can investigate the keys where information is stored for each item'''
dat['items'][0].keys()
```

Output:
```
dict_keys(['kind', 'id', 'etag', 'selfLink', 'volumeInfo', 'saleInfo', 'accessInfo', 'searchInfo'])
```

```python
# You can start building pretty long lines of code to access information deep 
#  in the structure. 
# Print out the ISBN_10 number for the book 
dat['items'][0]['volumeInfo']['industryIdentifiers'][0]['identifier']
```

Output:
```
'0553386794'
```

## Exercise 1 

Which of the Game of Thrones books is longest?

Get information about each book and print out the title and number of pages.  Then, report the book title and number of pages for the book with that is the longest.  

*Note, the API may return multiple entries for each isbn.  You may use the first entry for information.  If the information is missing a page number it is likely an audiobook, and you should then use the next entry for information.  If no entry has the title and page number information return the title as "no title" and the number of pages as '-1'.*

```python
''' Following is the isbn codes for Game of Thrones books. '''

isbns = ['0553386794', '0553579908', '9780345543981', '9780553582024', '9780553582017']
```

```python
'''
Iterate for each isbns to finds titles and pages for each item. 
Collect this information in a list. 
You can use "volumeInfo" to gather the information needed.
Print the title + the number of pages in the loop. 

Outside the loop:
- Convert the list to a DataFrame, column names 'Title' and 'NumPages' 
- Report the longest book in longestBookTitle and longestBookNumPages.
'''
ex1list = []

for i in isbns: 
    params = {'q': 'isbn:' + i}
    resp = requests.get(url, params)
    dat = resp.json()
    title = 'no title'
    pages = -1
    for i in range(0, len(dat['items'])):
      if 'title' in dat['items'][i]['volumeInfo']:
        title = dat['items'][i]['volumeInfo']['title']
      if 'pageCount' in dat['items'][i]['volumeInfo']:
        pages = dat['items'][i]['volumeInfo']['pageCount']
    ex1list.append((title, pages))
    print(title + " has " + str(pages) + " pages.")

titles = []
pageNums = []
for i in ex1list:
  titles.append(i[0])
for i in ex1list:
  pageNums.append(i[1])
ex1df = pd.DataFrame(titles, columns = ['Title'])
ex1df['NumPages'] = pageNums

longestBookTitle = ex1df['Title'].max()
longestBookNumPages = ex1df['NumPages'].max()
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-c67367e6bcd2> in <module>
     11 ex1list = []
     12 
---> 13 for i in isbns:
     14     params = {'q': 'isbn:' + i}
     15     resp = requests.get(url, params)

NameError: name 'isbns' is not defined
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Example: Government API - Iceland

Many cities or countries have begun making data available for developers and researchers. 
Iceland has created a single [API](http://docs.apis.is/) that has many endpoints including weather data, concerts, bus, earthquakes, bicycle counters, etc. 

Try out a few different endpoints and gather some data.  *Note, some of the endpoints are not available at this time, you may get a 404 or 500 error*

**Earthquake Information**

```python
resp = requests.get('https://apis.is/earthquake/is')
```

```python
resp
```

Output:
```
<Response [200]>
```

```python
resp.json()
```

Output:
```
{'results': [{'timestamp': '2017-10-13T12:07:24.000Z',
   'latitude': 63.976,
   'longitude': -21.949,
   'depth': 1.1,
   'size': 0.6,
   'quality': 58.73,
   'humanReadableLocation': '6,1 km SV af Helgafelli'},
  {'timestamp': '2017-10-13T09:50:50.000Z',
   'latitude': 65.124,
   'longitude': -16.288,
   'depth': 7.2,
   'size': 0.9,
   'quality': 78.51,
   'humanReadableLocation': '6,1 km NA af Herðubreiðartöglum'},
  {'timestamp': '2017-10-13T09:41:09.000Z',
   'latitude': 63.945,
   'longitude': -21.143,
   'depth': 7.4,
   'size': 0.2,
   'quality': 33.12,
   'humanReadableLocation': '6,5 km SSA af Hveragerði'},
  {'timestamp': '2017-10-13T09:37:45.000Z',
   'latitude': 65.114,
   'longitude': -16.3,
   'depth': 6.3,
   'size': 1.2,
   'quality': 90.01,
   'humanReadableLocation': '5,0 km NA af Herðubreiðartöglum'},
  {'timestamp': '2017-10-13T09:37:21.000Z',
   'latitude': 65.113,
   'longitude': -16.301,
   'depth': 5.9,
   'size': 1.4,
   'quality': 90.01,
   'humanReadableLoc
```

Under the `results` key there is a list of items with the earthquake information. 

Let's try to get this list in a DataFrame.

```python
# The JSON was already stored in "resp" and we only want the list of results 
#  under the key "results".  Therefore, we can take this information and 
#  re-serialize the information and let pandas read_json parse in the data 
eq = pd.read_json(json.dumps(resp.json()['results']))
eq.head()
```

Output:
```
timestamp  latitude  longitude  depth  size  quality  \
0 2017-10-13 12:07:24+00:00    63.976    -21.949    1.1   0.6    58.73   
1 2017-10-13 09:50:50+00:00    65.124    -16.288    7.2   0.9    78.51   
2 2017-10-13 09:41:09+00:00    63.945    -21.143    7.4   0.2    33.12   
3 2017-10-13 09:37:45+00:00    65.114    -16.300    6.3   1.2    90.01   
4 2017-10-13 09:37:21+00:00    65.113    -16.301    5.9   1.4    90.01   

             humanReadableLocation  
0          6,1 km SV af Helgafelli  
1  6,1 km NA af Herðubreiðartöglum  
2         6,5 km SSA af Hveragerði  
3  5,0 km NA af Herðubreiðartöglum  
4  4,9 km NA af Herðubreiðartöglum
```

If you get a "Value Error: protocol not known", then try wrapping the `json.dumps(resp.json()['results'])` in a method to write to a string using `StringIO` 

```python
from io import StringIO
eq = pd.read_json(StringIO(json.dumps(resp.json()['results'])))
```

**Football**

We can get information on football events. 



```python
resp = requests.get('https://apis.is/sports/football')
resp
```

Output:
```
<Response [200]>
```

```python
resp.json()
```

Output:
```
{'results': []}
```

Looks like no information listed for upcoming football matches. 

**Currency** 

```python
params = {'source' : 'lb'}
resp = requests.get('https://apis.is/currency', params)
resp
```

Output:
```
<Response [200]>
```

```python
resp.json()
```

Output:
```
{'results': [{'shortName': 'ISK',
   'longName': 'Íslensk króna',
   'value': 1,
   'askValue': 1,
   'bidValue': 1,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'USD',
   'longName': 'Bandarískur dalur',
   'value': 144.595,
   'askValue': 145.1,
   'bidValue': 144.09,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'GBP',
   'longName': 'Sterlingspund',
   'value': 161.245,
   'askValue': 161.81,
   'bidValue': 160.68,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'EUR',
   'longName': 'Evra',
   'value': 140.9,
   'askValue': 141.39,
   'bidValue': 140.41,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'CAD',
   'longName': 'Kanadískur dalur',
   'value': 105.19,
   'askValue': 105.56,
   'bidValue': 104.82,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'DKK',
   'longName': 'Dönsk króna',
   'value': 18.942,
   'askValue': 19.008,
   'bidValue': 18.876,
   'changeCur': 0,
   'changePer': '0.00'},
  {'shortName': 'N
```

**Weather Observations**

```python
params = {'stations': '1'}
resp = requests.get('https://apis.is/weather/observations/en', params)
resp

```

Output:
```
<Response [200]>
```

```python
resp.json()
```

Output:
```
{'results': [{'name': 'Reykjavík',
   'time': '2022-10-11 15:00:00',
   'err': '',
   'link': 'http://en.vedur.is/weather/observations/areas/reykjavik/#group=100&station=1',
   'F': '3',
   'FX': '8',
   'FG': '12',
   'D': 'W',
   'T': '6.0',
   'W': 'Light drizzle',
   'V': '10',
   'N': '100',
   'P': '993',
   'RH': '95',
   'SNC': '',
   'SND': '',
   'SED': '',
   'RTE': '',
   'TD': '5.3',
   'R': '4.5',
   'id': '1',
   'valid': '1'}]}
```

A description on what the codings stand for is available in the API documentation:  
http://docs.apis.is/#endpoint-weather

<!-- BEGIN QUESTION -->

## Exercise 2 

Look at forecasts for the Reykjavík weather station (station : 1).  Report the mean temperature and mean wind speed for the next 24 forecasts. 

```python
# Use the Iceland API to collect weather forecasts for the Reykjavik station. 
# Report the mean temperature and mean wind speed for the next 24 forecasts. 

params = {'stations': '1'}
resp = requests.get('https://apis.is/weather/forecasts/en', params)
dat = resp.json()
temps = []
speeds = []

for i in range(0, 24):
  temps.append(int(dat['results'][0]['forecast'][i]['T']))
  speeds.append(int(dat['results'][0]['forecast'][i]['F']))
meanTemp = sum(temps) / len(temps)
meanWS = sum(speeds) / len(speeds)
print("Mean Temperature (deg C): ", meanTemp)
print("Mean Wind Speed (m/s)   : ", meanWS)
```

Output:
```
Mean Temperature (deg C):  5.375
Mean Wind Speed (m/s)   :  4.666666666666667
```

<!-- END QUESTION -->

## Example: iTunes Content 

Apple has a simple [API](https://affiliate.itunes.apple.com/resources/documentation/itunes-store-web-service-search-api/) for looking up iTunes content.

```python
# api-endpoint
url = 'https://itunes.apple.com/search'

# For example let's search for lord of the rings ebooks 
params = {'term': 'lord+of+the+rings', 'entity': 'ebook' }

resp = requests.get(url, params)
```

```python
resp
```

Output:
```
<Response [200]>
```

```python
print(resp.text)
```

Output:
```
{
 "resultCount":2,
 "results": [
{
"artworkUrl60":"https://is1-ssl.mzstatic.com/image/thumb/Publication111/v4/bf/86/6b/bf866b79-4113-182f-a3f8-3cbde16234af/9780061798252_marketingimage.jpg/60x60bb.jpg", 
"artworkUrl100":"https://is1-ssl.mzstatic.com/image/thumb/Publication111/v4/bf/86/6b/bf866b79-4113-182f-a3f8-3cbde16234af/9780061798252_marketingimage.jpg/100x100bb.jpg", "artistViewUrl":"https://books.apple.com/us/artist/bernard-cornwell/2036463?uo=4", "artistIds":[2036463], "artistId":2036463, "artistName":"Bernard Cornwell", "genres":["Historical", "Books", "Fiction & Literature", "Action & Adventure"], "price":7.99, 
"description":"The&#xa0;fourth installment of Bernard Cornwell’s New York Times&#xa0;bestselling series chronicling the epic saga of the making of England, “like Game of Thrones, but real” (The Observer, London)—the basis for The Last Kingdom, the hit&#xa0;television series.<br /><br />The year is 885, and England is at peace, divided between the Danish kingdom to the
```

<!-- BEGIN QUESTION -->

## Exercise 3

Search for the 50 "The Expanse" e-books (search may return fewer). Create a data frame from the responses containing the `trackName`, `track ID`, `price`, and `averageUserRating`. Sort the results from highest to lowest price.

```python
url = 'https://itunes.apple.com/search'

# """ For example let's search for "The Expanse" ebooks """

params = {'term': 'expanse', 'entity': 'ebook' }
resp = requests.get(url, params) 

resp.json()
```

Output:
```
{'resultCount': 50,
 'results': [{'artworkUrl60': 'https://is3-ssl.mzstatic.com/image/thumb/Publication124/v4/1d/08/2b/1d082b19-5f04-4f29-c2e8-e43265b4cb25/9780316134675.jpg/60x60bb.jpg',
   'artworkUrl100': 'https://is3-ssl.mzstatic.com/image/thumb/Publication124/v4/1d/08/2b/1d082b19-5f04-4f29-c2e8-e43265b4cb25/9780316134675.jpg/100x100bb.jpg',
   'artistViewUrl': 'https://books.apple.com/us/artist/james-s-a-corey/433411981?uo=4',
   'trackCensoredName': 'Leviathan Wakes',
   'fileSizeBytes': 3645609,
   'formattedPrice': '$9.99',
   'trackViewUrl': 'https://books.apple.com/us/book/leviathan-wakes/id395522188?uo=4',
   'artistIds': [433411981],
   'description': '<b>From a <i>New York Times</i> bestselling and Hugo award-winning author comes a modern masterwork of science fiction, introducing a captain, his crew, and a detective as they unravel a horrifying solar system wide conspiracy that begins with a single missing girl. Now a Prime Original series.&#xa0;</b><br /><br />Humanity h
```

```python
obj = json.loads(resp.text)
# Alternatively 
obj2 = resp.json()
```

Try using at least two approaches to create the DataFrame, e.g., 

* *Method 1* - Keep track of rows in a list, convert nested lists to DataFrame.  Note, do not create an empty DataFrame and append entries in an iterator (this is not scalable)  
https://stackoverflow.com/questions/13784192/creating-an-empty-pandas-dataframe-and-then-filling-it/41529411#41529411
* *Method 2* - Use pandas `read_json` function to convert JSON to pandas object
* *Method 3* - Use `json_normalize` function that normalizes a semi-structured JSON data into a flat table.   
https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.io.json.json_normalize.html

```python
# Method 1 - Capture items in a list, convert to DataFrame
#Search for the 50 "The Expanse" e-books (search may return fewer). 
#Create a data frame from the responses containing the trackName, track ID, price, and averageUserRating. Sort the results from highest to lowest price.
rows = []
for i in range(4):
  rows.append([])
for i in range(50):
  rows[0].append(obj2['results'][i]['trackName'])
  rows[1].append(obj2['results'][i]['trackId'])
  rows[2].append(obj2['results'][i]['price'])
  if 'averageUserRating' in obj2['results'][i]:
    rows[3].append(obj2['results'][i]['averageUserRating'])
  else:
    rows[3].append(float("nan"))


q3df1 = pd.DataFrame({'trackName': rows[0], 'trackId': rows[1], 'price': rows[2]})

q3df1.shape
```

Output:
```
(50, 3)
```

```python
q3df1.head(15)
```

Output:
```
trackName     trackId  price
0                   Leviathan Wakes   395522188   9.99
1                     Caliban's War   469995594   9.99
2                    Abaddon's Gate   576438197   9.99
3                       Cibola Burn   721822343   9.99
4                     Nemesis Games   926360337   9.99
5                   Babylon's Ashes  1063592195  11.99
6                 Persepolis Rising  1215092254  11.99
7                    Tiamat's Wrath  1367091224  11.99
8                   Leviathan Falls  1557194850  14.99
9                         The Churn   834032860   2.99
10                     Gods of Risk   537353663   2.99
11      A Very Large Expanse of Sea  1330481470  11.99
12  The Butcher of Anderson Station   441318374   3.99
13                  The Vital Abyss  1042158204   3.99
14                     Strange Dogs  1227376539   2.99
```

```python
# Method 2 - Use pandas read_json function to 
# You may need to use StringIO function
# from io import StringIO  # Already imported at top of file 
file = StringIO(resp.text)
df = pd.read_json(file)
trackName = []
trackId = []
price = []
for res in df['results']:
  trackName.append(res['trackName'])
  trackId.append(res['trackId'])
  price.append(res['price'])

q3df2 = pd.DataFrame({'trackName':trackName, 'trackId':trackId, 'price':price})

q3df2.shape
```

Output:
```
(50, 3)
```

```python
q3df2.head(15)
```

Output:
```
trackName     trackId  price
0                   Leviathan Wakes   395522188   9.99
1                     Caliban's War   469995594   9.99
2                    Abaddon's Gate   576438197   9.99
3                       Cibola Burn   721822343   9.99
4                     Nemesis Games   926360337   9.99
5                   Babylon's Ashes  1063592195  11.99
6                 Persepolis Rising  1215092254  11.99
7                    Tiamat's Wrath  1367091224  11.99
8                   Leviathan Falls  1557194850  14.99
9                         The Churn   834032860   2.99
10                     Gods of Risk   537353663   2.99
11      A Very Large Expanse of Sea  1330481470  11.99
12  The Butcher of Anderson Station   441318374   3.99
13                  The Vital Abyss  1042158204   3.99
14                     Strange Dogs  1227376539   2.99
```

```python
obj2
```

Output:
```
{'resultCount': 50,
 'results': [{'artworkUrl60': 'https://is3-ssl.mzstatic.com/image/thumb/Publication124/v4/1d/08/2b/1d082b19-5f04-4f29-c2e8-e43265b4cb25/9780316134675.jpg/60x60bb.jpg',
   'artworkUrl100': 'https://is3-ssl.mzstatic.com/image/thumb/Publication124/v4/1d/08/2b/1d082b19-5f04-4f29-c2e8-e43265b4cb25/9780316134675.jpg/100x100bb.jpg',
   'artistViewUrl': 'https://books.apple.com/us/artist/james-s-a-corey/433411981?uo=4',
   'trackCensoredName': 'Leviathan Wakes',
   'fileSizeBytes': 3645609,
   'formattedPrice': '$9.99',
   'trackViewUrl': 'https://books.apple.com/us/book/leviathan-wakes/id395522188?uo=4',
   'artistIds': [433411981],
   'description': '<b>From a <i>New York Times</i> bestselling and Hugo award-winning author comes a modern masterwork of science fiction, introducing a captain, his crew, and a detective as they unravel a horrifying solar system wide conspiracy that begins with a single missing girl. Now a Prime Original series.&#xa0;</b><br /><br />Humanity h
```

```python
# Method 3 - Use json_normalize function
# The json_normalize function normalizes a semi-structured JSON data object into a flat table. 
# https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.json_normalize.html

tab = pd.json_normalize(obj)
tab = pd.json_normalize(tab['results'])
trackName = []
trackId = []
price = []
for col in tab:
  trackName.append(tab[col][0]['trackName'])
  trackId.append(tab[col][0]['trackId'])
  price.append(tab[col][0]['price'])

q3df3 = pd.DataFrame({'trackName':trackName, 'trackId':trackId, 'price':price})

q3df3.shape
```

Output:
```
(50, 3)
```

```python
q3df3.head(15)
```

Output:
```
trackName     trackId  price
0                   Leviathan Wakes   395522188   9.99
1                     Caliban's War   469995594   9.99
2                    Abaddon's Gate   576438197   9.99
3                       Cibola Burn   721822343   9.99
4                     Nemesis Games   926360337   9.99
5                   Babylon's Ashes  1063592195  11.99
6                 Persepolis Rising  1215092254  11.99
7                    Tiamat's Wrath  1367091224  11.99
8                   Leviathan Falls  1557194850  14.99
9                         The Churn   834032860   2.99
10                     Gods of Risk   537353663   2.99
11      A Very Large Expanse of Sea  1330481470  11.99
12  The Butcher of Anderson Station   441318374   3.99
13                  The Vital Abyss  1042158204   3.99
14                     Strange Dogs  1227376539   2.99
```

<!-- END QUESTION -->

## Example: TV Shows 

Here we can use an API on tv show information:  
http://api.tvmaze.com/

```python
# We can find the tvmaze id for a show based on the IMDB id. 
id_got = 'tt3032476'
resp = requests.get('http://api.tvmaze.com/lookup/shows?imdb=tt3032476')
```

```python
resp.json()
```

Output:
```
{'id': 618,
 'url': 'https://www.tvmaze.com/shows/618/better-call-saul',
 'name': 'Better Call Saul',
 'type': 'Scripted',
 'language': 'English',
 'genres': ['Drama', 'Crime', 'Legal'],
 'status': 'Ended',
 'runtime': 60,
 'averageRuntime': 64,
 'premiered': '2015-02-08',
 'ended': '2022-08-15',
 'officialSite': 'https://www.amc.com/shows/better-call-saul--1002228',
 'schedule': {'time': '21:00', 'days': ['Monday']},
 'rating': {'average': 8.6},
 'weight': 99,
 'network': {'id': 20,
  'name': 'AMC',
  'country': {'name': 'United States',
   'code': 'US',
   'timezone': 'America/New_York'},
  'officialSite': None},
 'webChannel': None,
 'dvdCountry': None,
 'externals': {'tvrage': 37780, 'thetvdb': 273181, 'imdb': 'tt3032476'},
 'image': {'medium': 'https://static.tvmaze.com/uploads/images/medium_portrait/399/998743.jpg',
  'original': 'https://static.tvmaze.com/uploads/images/original_untouched/399/998743.jpg'},
 'summary': '<p><b>Better Call Saul</b> is the prequel to the award-winni
```

We now know the TVmaze ID for Better Call Saul is **618**. 

## Exercise 4 

Calculate and report the min, mean, and max running time for Better Call Saul episodes by season in a DataFrame. 
Rows are indexed by season number (e.g., 1, 2, 3, ...) and Columns are "Min", "Mean", and "Max". 

Suggestion: 

* Create DataFrame of episodes information 
* Consider using the endpoint - http://www.tvmaze.com/api#show-episode-list
* Then use .groupby function to group by season
* Construct final DataFrame "bcs" to contain the requested information

```python
# Create a DataFrame "bcs" with row index of season number and 
#  columns of "Min", "Mean" and "Max" running time of the episodes for that season.

ep = requests.get('http://api.tvmaze.com/shows/618/episodes')
ep = ep.json()

seasons = []
runtimes = []

for e in ep:
  seasons.append(e['season'])
  runtimes.append(e['runtime'])

df = pd.DataFrame({'season':seasons, 'runtimes':runtimes})

means = df.groupby('season').mean()
maxs = df.groupby('season').max()
mins = df.groupby('season').min()

bcs = pd.DataFrame({'Min':mins['runtimes'], 'Mean':means['runtimes'],'Max':max['runtimes']})
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

Tests for q1 (points: 2)
```
>>> ex1df.shape[0] == 5 and ex1df.shape[1] == 2
True
```


Tests for q4 (points: 2)
```
>>> bcs.shape[0] == 6 and bcs.shape[1] == 3
True
```
