# A0 - Python Assignment 

Upon successful completion of this assignment, a student will be able to:

* Correctly setup Python environment on Campus Linux Machines
* Gain experience in formatting text using Markdown
* Submit assignment to Gradescope.

# Q1 - Setup

The following code looks to see whether your notebook is run on Gradescope (GS), Colab (COLAB), or the linux Python environment you were asked to setup.

```python
import re 
import os
import platform 
import sys 

# flag if notebook is running on Gradescope 
if re.search(r'amzn', platform.uname().release): 
    GS = True
else: 
    GS = False

# flag if notebook is running on Colaboratory 
try:
  import google.colab
  COLAB = True
except:
  COLAB = False

# flag if running on Linux lab machines. 
cname = platform.uname().node
if re.search(r'(guardian|colossus|c28)', cname):
    LLM = True 
else: 
    LLM = False

print("System: GS - %s, COLAB - %s, LLM - %s" % (GS, COLAB, LLM))
```

Output:
```
System: GS - False, COLAB - False, LLM - True
```

## Python Setup 

Check to make sure the correct version of Python was run. 

```python
pver = sys.version 
print(pver) 
```

Output:
```
3.8.10 (default, Jun  4 2021, 15:09:15) 
[GCC 7.5.0]
```

## Notebook Setup 

It is good practice to list all imports needed at the top of the notebook. You can import modules in later cells as needed, but listing them at the top clearly shows all which are needed to be available / installed.

If you are doing development on Colab, the otter-grader package is not available, so you will need to install it with pip (uncomment the cell directly below).

```python
# Only uncomment if you developing on Colab 
# if COLAB == True: 
#     print("Installing otter:")
#     !pip install otter-grader==4.2.0 
```

```python
# Import standard DS packages 
import pandas as pd 
import numpy as np
import matplotlib as mpl 
import matplotlib.pyplot as plt 
%matplotlib inline 

# Package for Autograder 
import otter 
grader = otter.Notebook()
```

## Environment Setup 

The python environment that is running is: 

```python
env1 = sys.executable
print(env1)
```

Output:
```
/home/campus24/dobidov/.conda/envs/cs4821/bin/python
```

```python
env2 =!conda info | grep 'active env'
print(env2)
```

Output:
```
['     active environment : cs4821', '    active env location : /home/campus24/dobidov/.conda/envs/cs4821']
```

Make sure that the environment you set up for the class is what is being used to execute your notebook. For example, the default name should be "cs4821". 

Next, we are going to look at all the packages installed. 

```python
import importlib
from packaging.version import Version, parse 

OK = '\x1b[42m[ OK ]\x1b[0m'
FAIL = "\x1b[41m[FAIL]\x1b[0m"

def import_version(pkg, req_ver, fail_msg=""):
    mod = None
    try:
        mod = importlib.import_module(pkg)
        ver = mod.__version__
        if Version(ver) != req_ver:
            print(FAIL, "%s version %s required, but %s installed."
                  % (lib, req_ver, ver))
        else:
            print(OK, '%s version %s' % (pkg, ver))
    except ImportError:
        print(FAIL, '%s not installed. %s' % (pkg, fail_msg))
        ver = ''
    return (mod, Version(ver), req_ver)

requirements = {'numpy': parse("1.21.6"), 'scipy': parse("1.9.3"),
                'matplotlib': parse("3.2.2"), 'pandas': parse("1.3.5"),
                'IPython': parse("8.7.0"), 'seaborn': parse('0.12.2'),
                'plotly': parse("5.9.0"), 'dill': parse('0.3.6'),
                'sklearn': parse("1.0.2")
                }

pks = []
for lib, required_version in list(requirements.items()):
    pks.append(import_version(lib, required_version))
```

Output:
```
[42m[ OK ][0m numpy version 1.21.6
[42m[ OK ][0m scipy version 1.9.3
[42m[ OK ][0m matplotlib version 3.2.2
[42m[ OK ][0m pandas version 1.3.5
[42m[ OK ][0m IPython version 8.7.0
[42m[ OK ][0m seaborn version 0.12.2
[42m[ OK ][0m plotly version 5.9.0
[42m[ OK ][0m dill version 0.3.6
[42m[ OK ][0m sklearn version 1.0.2
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

# Q2 

Enter the following formatted text into the Markdown cell below. 

In the same Markdown cell: 
* Add a first level heading with "a0-Python" 
* Add a second level heading with the due date formatted as:  `Due: %a %b %-d @ %-I:%-M%p`  
  Use `strftime()` and `strptime` format codes 
* Add boldface course "CS 4821" or "CS 5841"  
* On the next line, add the semester "Spring 2023" in italics 
* Copy over the course objectives from the [syllabus](https://mtu.instructure.com/courses/1441806/assignments/syllabus) as an unordered list, e.g., "Upon successful completion of this course, ..." 
* Create a table with the grading scale A-F from the syllabus
* Add a link to the Office of Academic and Community Conduct Resource Center for Students 

# a0-Python

## Due: Thu Jan 26 @ 11:59PM

**CS 5841**

*Spring 2023*

Upon successful completion of this course, students will be able to:

-  Describe advanced data mining concepts and techniques
-  Apply the techniques of clustering, classification, association finding, and visualization on sample and real-world data sets
-  Select and apply data mining tools to real-world problems
-  Formulate a real-world data mining problem: find, get, and clean the data, select and implement appropriate methods, evaluate the performance of the methods, etc.
-  Discuss and communicate the results and methods of applying data mining techniques to a problem

| Percentage | Letter Grade |
|:-----------|:-------------|
| > 93% |A|
| [88 - 93%) | AB |
| [83 - 88%) | B |
| [78 - 83%) | BC |
| [73 - 78%) | C |
| [68 - 73%) | CD |
| [60 - 68%) | D |
| < 60% | F |

[Office of Academic and Community Conduct Resource Center for Students](https://www.mtu.edu/conduct/)

<!-- END QUESTION -->

# Checks before submitting 

Make sure to run this notebook on the campus linux machines and submit to Gradescope the zip file that is created. 

**FAQ:** 

* Q: I notice my zip file is growing as I repeatedly run the notebook? 
* A: This assignment is graded from a log file called `.OTTER_LOG`.  Other assignments will be graded by rerunning your code. Before rerunning the entire notebook `Cell -> Run All`, delete the `.OTTER_LOG` file.  



* Q: I am getting different package numbers? 
* A: Make sure you follow the instructions to set up the Python environment on the campus linux machines.  


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

**NOTE** the submission must be run on the campus linux machines.  See the instruction in the Canvas assignment.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
<IPython.core.display.HTML object>
```

```python

```

Tests for q1 (points: 13)
```
>>> GS == False 
True
```
```
>>> COLAB == False 
True
```
```
>>> LLM == True 
True
```
```
>>> pver[0:6] == '3.8.10' 
True
```
```
>>> pks[0][1] == parse('1.21.6')
True
```
```
>>> pks[1][1] == parse('1.9.3')
True
```
```
>>> pks[2][1] == parse('3.2.2')
True
```
```
>>> pks[3][1] == parse('1.3.5')
True
```
```
>>> pks[4][1] == parse('8.7.0')
True
```
```
>>> pks[5][1] == parse('0.12.2')
True
```
```
>>> pks[6][1] == parse('5.9.0')
True
```
```
>>> pks[7][1] == parse('0.3.6')
True
```
```
>>> pks[8][1] == parse('1.0.2')
True
```
